# SLU16 - Workflow: Exercise Notebook

```python
import random
import math
import pandas as pd
import matplotlib.pyplot as plt

from sklearn import linear_model
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import MinMaxScaler
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score
from sklearn.base import TransformerMixin
from sklearn.model_selection import train_test_split
from sklearn.impute import SimpleImputer

from hashlib import sha1 # just for grading purposes
import json # just for grading purposes

from utils import workflow_steps, data_analysis_steps
from utils import get_dataset
from utils import plot_confusion_matrix

def _hash(obj, salt='none'):
    if type(obj) is not str:
        obj = json.dumps(obj)
    to_encode = obj + salt
    return sha1(to_encode.encode()).hexdigest()
import warnings
warnings.filterwarnings('ignore')
```

In this notebook we will be covering the following:

* **1. Workflow**
 * Step 1: Get the data
 * Step 2: Data analysis and preparation
   * Step 2.1 Data analysis
   * Step 2.2 Dealing with data problems
   * Step 2.3 Feature engineering
   * Step 2.4 Feature selection
 * Step 3: Model training
 * Step 4: Evaluate results
* **2. Pipelines and Custom Objects**
    * Step 2.1: Pipelines 
        * Step 2.1.1: Doing it "the hard way"
        * Step 2.1.2: What is a pipeline
        * Step 2.1.3: Setting up a pipeline
    * Step 2.2: Custom Objects
        * Step 2.2.1: Custom Transformers
        * Step 2.2.2: Custom Estimators

<div class="alert alert-info">
    A <b>data science workflow</b> defines the phases (or steps) in a data science project. Using a well-defined data science workflow is useful not only to you, but your teammates as it provides a simple way to clearly structure and organize a data science project. Across this specialization we've been covering the different steps in this workflow, but how well are you familiarised with them?
</div>

## Exercise 1: Workflow

### Exercise 1.1 - Overall workflow steps

What are the basic workflow steps?

You probably know them already, but we want you to really internalize them. We've given you a list of steps in `workflow_steps`, but it appears that, not only does it have too many steps, some are _probably_ wrong, as well.

Select the correct ones and reorder them!

```python
print("Workflow steps:")
for i in range(len(workflow_steps)):
    print(i+1, ': ', workflow_steps[i])
```

<!-- grade_id: cell-eb636145aa25eef4, points: 0, grade: false, solution: true -->
```python
# Exercise 1.1. Filter and sort the names of the steps in the workflow_steps list
# workflow_steps_answer = [...]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c9ae302d2c9c48c, points: 1, grade: true, solution: false -->
```python
assert _hash(len(workflow_steps_answer), 'salt0') == '253bd16d7f97f7d7295d404825e32985d257e11f', "your workflow size doesn't look right! Don't forget to remove steps that shouldn't be there"
assert _hash([step.lower() for step in workflow_steps_answer], 'salt0') == '701e2306da9bfde36382bdb6feb80a354916ebf4', "your workflow order doesn't look right! Some steps might be out of place"

```

## Exercise 1.2: Data Analysis and Preparation

There are way too many substeps in the **Data Analysis and Preparation** step to group them all under a single category. We've given you another list of steps: `data_analysis_steps`.

Aside from being shuffled, it should be fine but keep an eye out. You never know what to expect...

```python
print("Data Analysis and Preparation steps:")
for i in range(len(data_analysis_steps)):
    print(i+1, ': ', data_analysis_steps[i])
```

<!-- grade_id: cell-2dab17640925e0d9, points: 0, grade: false, solution: true -->
```python
# Exercise 1.2. Filter and sort the names of the steps in the data_analysis_steps list
# data_analysis_steps_answer = [...]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d6087c6110d2371, points: 1, grade: true, solution: false -->
```python
assert _hash(len(data_analysis_steps_answer), 'salt0') == '3aa298b947129ff226ba1e3f545f8b5ffaa6e557', "your workflow size doesn't look right! Don't forget to remove steps that shouldn't be there"
assert _hash([step.lower() for step in data_analysis_steps_answer], 'salt0') == '658ab90eff4a0cea2bfb51cc89c8db5b4121fa86', "your workflow order doesn't look right! Some steps might be out of place"

```

<img src="media/spanish_inquisition.gif" width="400" />

## Exercise 2 - Walking down the yellow (workflow) path

There is no template for solving a data science problem. The roadmap changes with every new dataset and new problem. But we do see similar steps in many different projects. Regardless, some steps are fairly common in any process. Let's go through them one by one.

### Exercise 2.1 - Objective

Every DS analysis should start with one question: **What is the problem you are trying to solve?** Clearly stating your problem is the first step to solving it and without a clear problem, you could find yourself down a data-science rabbit-hole.

For this workflow, we are going to analyze a **a randomly generated dataset**. The objective? 

<div class="alert alert-info">
    Build a model to predict the <b>value of y</b> given a set of features.
</div>

#### Exercise 2.1.1 - Objective
Let's start by importing the dataset

```python
df, y = get_dataset()  # preloaded dataset
df['y'] = y
df.head()

```

<div class="alert alert-warning">
⚠️ Is the objective clear to you?
This is just a yes or no question, no need for code here! :P 
</div>

<!-- grade_id: cell-7bac9e5a5b30e333, points: 0, grade: false, solution: true -->
```python
#answer_2_1_1=False
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dad9e0e457bf29a1, points: 0, grade: true, solution: false -->
```python
assert answer_2_1_1, "Don't make the panda sad!"
```

#### Exercise 2.1.2 - Clasifying the problem
Now that we have our data imported into Pandas and we've checked out the first few rows of our dataframe, there's a few questions we need to answer before we move on:

- *A*: Is this **supervised learning** or **unsupervised learning**? 
- *B*: Is this a **classification problem** or is it a **regression problem**? 
- *C*: Is this a **prediction problem** or an **inference problem**?

Keeping our **objective in mind** how would you classify this problem?

In our example, we are going to be using **regression** to **predict** value y from the set of features.

Save, in `answer_2_1_2`, the value from **A, B and C** that apply to our problem!

<!-- grade_id: cell-26745fdc9dca03d8, points: 0, grade: false, solution: true -->
```python
# Remove from the string what doesn't apply to our problem (including the '/')
#answer_2_1_2 = ["supervised/unsupervised learning","classification/regression problem", "prediction/inference problem"]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-80d1e9191be2cad6, points: 1, grade: true, solution: false -->
```python
assert _hash(len(answer_2_1_2), 'salt0') == 'bb6591f296df2e5dd0324184f4348795b760f11d', "your answer doesn't have the correct size. I've asked you to pick the correct option for three questions"
assert _hash([step.lower() for step in answer_2_1_2], 'salt0') == 'd5cf73be53ec62d0d1ca23376f66f5a30cce8f67', "One or more of your answers is incorrect."
```

### Exercise 2.2 - Data Exploration and Data Cleaning

Back to our data! Let’s determine which variable is our target and which features we think are important.
Our target is going to be the column titled **y** and our features are going to be the columns not containing the words **arm** or **leg** (assume we got this information from our boss or client). Remove all of the columns we don’t need for this analysis. Order the columns lexicographically.

<div class="alert alert-warning"> 
⚠️ <b>NOTE: </b>lexicographic sorting means basically that the language treats the variables as strings and compares character by character ("200" is greater than "19999" because '2' is greater than '1').
</div>
    
Save the resulting dataframe in `df_clean`. 

Remember, in this case we're telling you what columns are to be kept, but this is something that is part of your workflow process. **A good data exploration and data cleaning is a key factor in the outcome of your model!**

<!-- grade_id: cell-138aa6edcbd29190, points: 0, grade: false, solution: true -->
```python
#df_clean = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-797514f689b49b52, points: 2, grade: true, solution: false -->
```python
assert isinstance(df_clean, pd.DataFrame), "Should be a dataframe"
assert df_clean.shape == (100, 21), "The shape of the dataframe is different than expected. Have you dropped the uncessessary columns?"
assert _hash([step.lower() for step in df_clean.columns], 'salt0') == '64d0d2fe9259036f208c2ef818f158ee866a1968', "One or more of your column headers is incorrect."

```

### Exercise 2.3 - EDA

Exploratory data analysis (EDA) gives the data scientist an opportunity to really learn about the data they are working with. 

Throughout the EDA process, I clean the data. Data from the real world is *very messy*. As I work through the EDA process and learn about the data, I take notes on things I need to fix in order to conduct my analysis. Most times, **Data cleaning and EDA go hand in hand for me**.

The first things I check are data types. Getting all of the values in the correct format is important. This can involve stripping characters from strings, converting integers to floats, or many other things.

What type of variables are our features? 

<!-- grade_id: cell-4dfbeeaad55f42af, points: 0, grade: false, solution: true -->
```python
#answer_2_3 = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e82bcea511e8aa79, points: 2, grade: true, solution: false -->
```python
assert _hash(answer_2_3, 'salt0') == '70d410ba8247cfc821d4442f5ad0b48d983138e0', "Data type is incorrect."
```

### Exercise 2.4 - Impute missing values
Finding missing values is quite common. Just replace them in our clean dataframe with the mean of the corresponding column/feature.

<!-- grade_id: cell-fc4703b89bf0d85b, points: 0, grade: false, solution: true -->
```python
#df_clean = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8443173a170a44d8, points: 2, grade: true, solution: false -->
```python
assert isinstance(df_clean, pd.DataFrame), "Should be a dataframe"
assert df_clean.isna().sum().sum() == 0, "Missing values are still present"
assert df_clean.shape == (100, 21), "The shape of the dataframe is different than expected. Have you dropped the rows with missing data?"
assert math.isclose(df_clean.values.sum(), 127.511, abs_tol=0.01), "Are you repleacing the missing values by the mean of each column?" 
```

Uff! That took quite some time, but now we have a clean and tidy dataframe to work with!

## Exercise 3 - Baseline modeling

As a data scientist, you will build a lot of models. You will use a variety of algorithms to perform a wide variety of tasks. You will need to use intuition and experience to decide when certain models are appropriate! 

But when constructing your baseline model, the simpler, the better! 

Let's start!

### Exercise 3.1 - Separate your target value 

Separate into `X` and `y` your features and your target. Order the columns in `X` as in exercise 2.2. 

<!-- grade_id: cell-c1c072726b679fe8, points: 0, grade: false, solution: true -->
```python
#X = ...
#y = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d02011f7f3426b07, points: 1, grade: true, solution: false -->
```python
assert _hash(list(y), 'salt0') == '917a4dd4b91497ba8c83f4d64b5d27642af0ce4f', "Have you picked the right column as the target?"
assert X.shape == (100, 20), "The shape of the dataframe is different than expected. Have you dropped the target?"
assert _hash(sorted(X.columns), 'salt0') == 'df65a7cdfbc4e774e8e18dbc46c7faeb42dc4c27', "Have you included the right columns in X?"
```

### Exercise 3.2 - Split data

Split your dataset into test and train data, using `test_size=0.2` and `random_state=42` on your `train_test_split()`

<!-- grade_id: cell-0a59ceb29b6e7a7b, points: 0, grade: false, solution: true -->
```python
#X_train, X_test, y_train, y_test = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eda4965289c315c6, points: 1, grade: true, solution: false -->
```python
assert (X_train.shape, X_test.shape, y_train.shape, y_test.shape) == ((80, 20), (20, 20), (80,), (20,)), "have you split the data correctly? Test size should be 0.2"
```

### Exercise 3.3 - Scale your data
As we are not sure whether the features units of measures are in the same scale, you should scale your X_train and X_test dataframes separately, to avoid data leakage. Use `MinMaxScaler`.

<!-- grade_id: cell-fe58c296cca28681, points: 0, grade: false, solution: true -->
```python
# corr_X_train = ...
# corr_X_test = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python
sum(corr_X_test[:,4])
```

<!-- grade_id: cell-cc5e79e6d94f3521, points: 1, grade: true, solution: false -->
```python
assert math.isclose(sum(corr_X_train[:,1]), 42.14, abs_tol=0.01), "have you used the correct scaler?"
assert math.isclose(sum(corr_X_test[:,4]), 11.78, abs_tol=0.01), "have you used the correct scaler?"

```

### Exercise 3.4 - Finally! The model!

We can finally make our predictions with a simple Random Forest Classifier!

<!-- grade_id: cell-6423f554ffdba88d, points: 0, grade: false, solution: true -->
```python
#rf = RandomForestClassifier()
#model = rf.fit(..., ... )
#predictions = rf.predict(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-408ce12cfa728ad1, points: 2, grade: true, solution: false -->
```python
assert predictions.size == 20
plot_confusion_matrix(y_test, predictions)

print ('Accuracy score:', accuracy_score(y_test, predictions)) 
```

Our model is not performing bad at all! If you want to improve it, you should make **small alterations**, **one at a time**! Keeping track of your changes is crucial to know exactly what change is helping or hurting your model! 

## Exercise 4 - Pipelines!!!!! 

We've already loaded and splitted a dataset for the following exercises. They're stored in the `new_X_train`, `new_X_test`, `new_y_train` and `new_y_test` variables.

In a perfect world, where you have all your data clean and ready-to-go, you can create your pipeline with just Scikit-learn's Transformers. However, in the real world, that's not the case, and you'll need to create custom Transformers to get the job done. Take a look at the data set, what do you see?

```python
new_X, new_y = get_dataset()  # preloaded dataset
new_X_train, new_X_test, new_y_train, new_y_test = train_test_split(new_X, new_y, test_size=0.33)
```

```python
# use this cell to explore the data
# do you notice something interesting in the data?
```

As you can see, it's the exact same dataset from the previous exercises! The same issues have returned:

- There are 4 columns whose name starts with either arm or leg which are all filled with gibberish.
- There are some missing values in some columns.

So, first things first, let's get rid of those columns with gibberish through a **Custom Transformer**, so we can plug it in a Scikit Pipeline after.

### Exercise 4.1 - Custom Transformer

<!-- grade_id: cell-dc3fa4786483541e, points: 0, grade: false, solution: true -->
```python
# Create a pipeline step called RemoveLimbs that removes any
# column whose name starts with the string 'arm' or ´leg´

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b91664f72233c2c7, points: 3, grade: true, solution: false -->
```python
assert issubclass(RemoveLimbs, TransformerMixin)
assert _hash(sorted(RemoveLimbs().fit_transform(new_X).columns), 'salt5') == '71443dfc3077d773d4c74e958dadf91dc2cc148a'
assert _hash(list(map(lambda col: col.startswith('arm') or col.startswith('leg'), RemoveLimbs().fit_transform(new_X_train).columns)), 'salt6') == 'ce45cf3759d2210f2d1315f1673b18f34e3ac711'
```

### Exercise 4.2 - Pipelines are the best!

Now that we have our Custom Transformer in place, we can design our pipeline! 

Create a pipeline with the following steps:

1. Removes limbs columns
2. Imputes missing values with the mean
3. Has a Random Forest Classifier as the last step

You may use `make_pipeline` to create your pipeline with as many steps as you want as long as the first two are the Custom Transformer you developed previously, a `SimpleImputer` as the second step, and a `RandomForestClassifier` as the last step. Save your pipeline into a variable named `pipeline`!

<!-- grade_id: cell-178a0f61ec49c4ac, points: 0, grade: false, solution: true -->
```python
# pipeline = make_pipeline(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3939f2f56f82f742, points: 3, grade: true, solution: false -->
```python
assert _hash(pipeline.steps[0][0], 'salt7') == '471b02068ac2c4f479c2e9f85f4b3dc2179bb841'
assert _hash(pipeline.steps[1][0], 'salt8') == 'ca83eaea1a7e243fa5574cfa6f52831166ee0f32'
assert _hash(pipeline.steps[-1][0], 'salt9') == '0d66ba4309ad4939673169e74f87088dcadd510b'
```

Does it work? Let's check it out on our dataset!

```python
pipeline.fit(new_X_train, new_y_train)
new_y_pred = pipeline.predict(new_X_test)
accuracy_score(new_y_test, new_y_pred)
```

It doesn't get much simpler that this, does it?

For an extra challenge, go back to exercises 2 and 3 and follow our workflow but with a Pipeline! For each special processing we've done, you can create a custom transformer for that column! 

Dominating pipelines and custom transformers can be a huge time saver! And there's the hackathon ahead...

**Good luck!**