# BLU01 - Exercises Notebook

**Important Note**

Attention Windows users: The grader will run on Linux, and using power shell statements will not work.

```python
import chardet
import hashlib # for grading purposes
import math
import numpy as np
import csv
import pandas as pd
import random
```

## Q1: Use a shell command to count lines

Count the total number of lines of files data/exercises/dog.txt and data/exercises/dog_intelligence.txt and store the result in variable called count_total.  
There is more than one way to do this. Hint: concatenate both files.





<!-- grade_id: cell-05a3d2245d57305d, points: 0, grade: false, solution: true -->
```python
# count_total = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e69c6137e73d90c, points: 2, grade: true, solution: false -->
```python
expected_hash = '2c624232cdd221771294dfbb310aca000a0df6ac8b66b696d90ef06fdefb64a3'
assert hashlib.sha256(str(count_total).encode()).hexdigest() == expected_hash, "count_total is wrong"
```

![title](./media/dog_1.jpeg)


Now store in 2 variables the first fourth and the last fourth of the lines of the file data/exercises/dog.txt.

<!-- grade_id: cell-14cf2619322a8763, points: 0, grade: false, solution: true -->
```python
# NOT GRADED optional exercise!!
# Add the first and last fourth of the lines to the variables first_third and last_third. 
# Make it work to any files using count_total_dog / 4 for instance in the bash commands
# count_total_dog = 
# first_fourth = ...
# last_fourth = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-20d96111430874d4, points: 0, grade: true, solution: false -->
```python
# Test the first_fourth and last_fourth exercise by running the following (ungraded) asserts.
assert first_fourth[-1][0] == 'T'
assert last_fourth[0][0] == 'S'
```

## Q2: Read a file with specific delimiter

Read file **data/exercises/breed_info.csv** into a pandas DataFrame. This file contains information about dog breeds, including minimum and maximum heights and weights.

First, you should preview the file using a shell command in order to find out the used delimiter, and other properties of this file.

Then, you should use the function read_csv to read the data into a DataFrame. The resulting DataFrame should have the first column as index.

<!-- grade_id: cell-942e0590467e20d9, points: 0, grade: false, solution: true -->
```python
# Use a shell command to preview the data
# ! ...
# YOUR CODE HERE
raise NotImplementedError()

# Use function read_csv to read the data into a DataFrame
# df2 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c275c9acaa9c74d, points: 2, grade: true, solution: false -->
```python
assert df2.loc['Bernese Mountain Dog', 'height_low_inches'] == 23.0, "df2 data is wrong, the lower range of the height of the Bernese Mountain Dog is incorrect."
assert len(df2.columns) == 5, "df2 data is wrong, the number of columns does not match what is expected."
assert len(df2) == 148, "df2 data is wrong, the length is different than expected."
assert df2.index[34] == 'English Setter', "df2 data is wrong, the positions are wrong."

expected_hash = 'd1ac8a80937c65bd103bb3aab437b45ee947eff75901f671d68f895e35eb09d2'
assert hashlib.sha256(df2.index.name.encode()).hexdigest() == expected_hash
```

## Q3: Read a csv file with problems

Read file **data/exercises/breed_info_w_problems.csv** using function `read_csv`. Pay attention to the following:
* use the first column as index
* there are some inputs in the file that should be interpreted as NaN, make sure you select the right one when reading the file!

<!-- grade_id: cell-2771b707556c08d2, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/breed_info_w_problems.csv with read_csv
# df3 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-601b354802964776, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df3.loc['Poodle Standard','weight_high_lbs']), "df3 data is wrong, there seems to be inputs that were not interpreted as NaN."
assert df3.loc['Kuvasz', 'weight_low_lbs'] == 70, "df3  data is wrong, the lower range of the weight of the Kuvasz cannot be found."
assert df3.index[3] == 'Bloodhound', "df3 data is wrong, the positions of the object are wrong."

expected_hash = 'd1ac8a80937c65bd103bb3aab437b45ee947eff75901f671d68f895e35eb09d2'
assert hashlib.sha256(df3.index.name.encode()).hexdigest() == expected_hash
```

## Q4: Repair a csv file after importing
Read a new file **data/exercises/pet_stats.csv** using the function `csv.reader()`. This file includes information about the different pet populations. Be sure you don't miss any lines with relevant information! 

* use the csv module to import everything to a list of lists
* create a df with information from the first four columns (countries)
* format the columns with the right type (don't forget to change the date from a string)
* set the index of the df to the `country`

<!-- grade_id: cell-119041388e17fb72, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/microsoft_info_w_problems.csv using csv.reader() and add result to
# lines = ...
# YOUR CODE HERE
raise NotImplementedError()
# create a dataframe using the line list with only 7 columns (country and pet types)
# df4 = pd.DataFrame(...
# YOUR CODE HERE
raise NotImplementedError()
# replace empty values with nan (np.nan)
# YOUR CODE HERE
raise NotImplementedError()
# set types per dataframe column, including the country. Always use float64 when floats are needed.
# YOUR CODE HERE
raise NotImplementedError()
# set country as the new index of the dataframe
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a42299f9047590d5, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df4.loc['Estonia', 'bird_pop']), "df4 content is wrong"
assert df4.loc['Slovakia', 'cat_pop'] == 32000, "df4 content is wrong"

mean_cat_pop = df4['cat_pop'].mean()
assert math.isclose(6915800.0, mean_cat_pop, rel_tol=1e-3), "df4 content is wrong"

assert df4['bird_pop'].dtype == np.float64, "df4 types are wrong"
assert df4.index.dtype == np.dtype('object'), "df4 types are wrong"

expected_hash = 'aff64e4fd520bd185cb01adab98d2d20060f621c62d5cad5204712cfa2294ef7'
assert hashlib.sha256(df4.index.name.encode()).hexdigest() == expected_hash
```

## Q5: Read a JSON file

Read the file **data/exercises/dog_pop_country_2023.json**.

<!-- grade_id: cell-5c2125479f7e50e5, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/dog_pop_country_2023.json with read_json
# df5 = read_json(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f085fa2cdad9319, points: 2, grade: true, solution: false -->
```python
assert len(df5) == 37
assert set(df5.columns) == {'EstimatedDogPopulation', 'dogsPerPerson'}, "df5 columns are wrong"

expected_hash = 'fab1d9dc541cefd3541a1ca75af65d04d5eb3e035369c8ef7d50c4b12282d121'
assert hashlib.sha256(str(df5.loc[:,'dogsPerPerson'].sum()).encode()).hexdigest() == expected_hash, "df5 content is wrong"
```

## Q6: Read an Excel file

Read file **data/exercises/dog_breeds.xlsx** using function read_excel. This file has info on different dog breeds, namely behavioural and size characteristics. In this exercise, please do the following:

* you should grab the dat in sheet "breed_info"
* use column 'Name' as index
* make sure you keep only the rows and columns with data
* set the variable best_dogs_w_children with the number of dogs with the maximum score for 'good_with_children'

<!-- grade_id: cell-017a7c31b0ddc38e, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/dog_breeds.xlsx with read_excel
# df6 = read_excel(...)
# YOUR CODE HERE
raise NotImplementedError()

# best_dogs_w_children = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91d5f07e40585680, points: 2, grade: true, solution: false -->
```python
assert best_dogs_w_children == 42
assert isinstance(df6, pd.DataFrame)
expected_hash = 'f0a0278e4372459cca6159cd5e71cfee638302a7b9ca9b05c34181ac0a65ac5d'
assert hashlib.sha256(
        df6.loc['Hovawart', 'good_with_strangers']
    ).hexdigest() == expected_hash, "df6 is wrong"
```

## Q7: Find the encoding of a file

Find the encoding used in file **data/exercises/dogs_intelligence_sample.csv**, using the method that was shown in the Learning Units.

Then, read the data into a DataFrame, using the read_csv method and find the `Breed` with an upper limit of repetitions to understand new commands (`reps_upper`) of 15.

<!-- grade_id: cell-bd6944fd63bb38d8, points: 0, grade: false, solution: true -->
```python
# Find the encoding of file data/exercises/dogs_intelligence_sample.csv
# encoding = ...
# YOUR CODE HERE
raise NotImplementedError()
# Read the file into a DataFrame
# df7 = ...
# YOUR CODE HERE
raise NotImplementedError()

# Breed with reps_upper equal to 15
# city_found = .... (assign a string)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-687fb83c97e03355, points: 2, grade: true, solution: false -->
```python
assert isinstance(df7, pd.DataFrame)

expected_hash_1 = 'c543ece81605c7d202121c62080a0db4020fc2c75bfac35d101d7f3e93c93949'
assert hashlib.sha256(encoding.encode()).hexdigest() == expected_hash_1, "encoding is wrong"

expected_hash_2 = '0d9abf1f6b31fb0cdc875d352b47751ce91ca57075087365502ea648b3329a19'
assert hashlib.sha256(breed_found.encode()).hexdigest() == expected_hash_2, "breed_found is wrong"
```

## Q8: Import a Random Sample of a Big File

Consider the file **data/exercises/dog_intelligence.tsv**. Let's imagine this file is really huge, with a lot of rows!  Read the file using a random sample of 10 rows. Count the actual lines with `wc` 

<!-- grade_id: cell-fb39107093dd73fb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/dog_intelligence.tsv with wc and save the number of lines
# in lines_in_file (as an integer), using 
# lines_in_file = ...
# Notice that the header is not a line..

# YOUR CODE HERE
raise NotImplementedError()

# make parameter rows_to_skip equal to the lines you want to skip loading 
# don't forget: 10 rows should be fetched
# rows_to_skip = ...
# YOUR CODE HERE
raise NotImplementedError()
# Create a df8 dataframe with the sampled values
# don't forget: you want to keep the header plus 10 rows in the new dataframe.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec15a8be42137a01, points: 2, grade: true, solution: false -->
```python
# getting remaining list of dog breeds, removing rows_to_skip from the file
df_helper = pd.read_csv( 
    'data/exercises/dog_intelligence.tsv', 
    sep='\t',
    header=0 
)

total_indexes = list(df_helper.index)
for s in rows_to_skip:
    total_indexes.remove(s-1)

expected_hash = '13671077b66a29874a2578b5240319092ef2a1043228e433e9b006b5e53e7513'
assert hashlib.sha256(str(lines_in_file).encode()).hexdigest() == expected_hash, "lines_in_file are wrong"
assert isinstance(df8, pd.DataFrame)
assert list(df8['Breed'])==list(df_helper.iloc[total_indexes]['Breed'])
assert df8.shape[0]==11, "df8 size is wrong"

expected_hash = 'e7f6c011776e8db7cd330b54174fd76f7d0216b612387a5ffcfb81e6f0919683'
assert  hashlib.sha256(str(df8.shape[1]).encode()).hexdigest() == expected_hash

```

## Q9: Loading a Big File

Read file **data/exercises/dog_intelligence.tsv** using chunks keep only the columns `Breed`, `Classification`, and `obey`.

Note that:
* the file should be read in chunks of 20 breeds
* the missing values for `Classification` should be removed (filtered in each chunk)
* the `obey` should be converted to type float (in each chunk)
* the index should be incremental starting from 0 (i.e, you don't need to read any column as the index)

At the end, calculate the average `obey`

Tip: Be sure you investigate the data you are about to load before doing any code.

<!-- grade_id: cell-5106c25a705296e8, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/dog_intelligence.tsv
# the chunks should be appended in a list called chunk_arr
# YOUR CODE HERE
raise NotImplementedError()

# df9 should be the final dataframe with concatenated chunks
# Resulting average should go on lit_avg variable
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f60a22c3465d5b83, points: 2, grade: true, solution: false -->
```python
expected_hash = 'f66229d92035c72b54407f3bc882eaadff7e3b97969f785c4b5caf4bca46e84f'
assert hashlib.sha256(str([len(c) for c in chunk_arr]).encode()).hexdigest() == expected_hash, "error on chunk_arr sizes, did you correctly remove the missing values?"

assert df9.loc[df9['Breed']=='Basenji', 'obey'].values[0] == 10.0, "df9 data is wrong"
assert df9.dtypes['Breed'] == object, "df9 structure is wrong, check data types"
assert df9.dtypes['obey'] == float, "df9 structure is wrong, check data types" 

assert math.isclose(56.43, lit_avg, rel_tol=1e-3), "lit_avg is wrong"
```

## Q10: Calculate average values of `reps_lower` using chunks and avoiding a complete data frame in memory.

Using chunks, read file **data/exercises/dog_intelligence.tsv**, avoid incomplete rows within `reps_lower` and calculate the average of `reps_lower` without loading all data simultaneously. 

Use a similar approach of the previous question but don't create any dataframe neither any list with chunks; ***Hint: Use the average definition***



<!-- grade_id: cell-57fd7700885246eb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/dog_intelligence.tsv
# the final average should be in the variable final_avg
# You should increment 2 variables in each chunk and use them at the end to calculate final_avg, call them rep_a, rep_b
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-0edcf1260d7886c9, points: 2, grade: true, solution: false -->
```python
assert math.isclose(25.58823529411765, final_avg, rel_tol=1e-1), "final_avg is wrong"
assert int(rep_b) == 3480 or int(rep_a)==136, "rep_a or rep_b is wrong is wrong"
```

Congrats, you were able to deal with all the messy data of this notebook!

![title](./media/dog_2.jpeg)