# SLU10: Metrics for Classification -- Exercises

In this notebook, we'll practice all the classification metrics you've learned. This is the roadmap:

- Understanding the problem with accuracy
- Understanding FP, TP, FN, TN
- Understanding precision and recall
- Understanding the ROC Curve 
- Understanding AUROC
- Which model is better for a particular circumstance
- Using these metrics in day to day 

```python
import pandas as pd
import math 
import numpy as np
from utils import plot_confusion_matrix, hash_answer
from matplotlib import pyplot as plt 
```

```python
# Run this cell - we will use this function in the exercises.

def load_data():
    # Loads classifier probabilities dataframe
    df = pd.read_csv("data/classifier_prediction_scores.csv")
    return df
```

## Useful information

In this notebook, we will work with a dataset with probabilities and class labels like this:

<table border="1" class="dataframe">  <thead>    <tr style="text-align: right;">      <th></th>      <th>probas</th>      <th>target</th>    </tr>  </thead>  <tbody>    <tr>      <th>0</th>      <td>0.288467</td>      <td>0</td>    </tr>    <tr>      <th>1</th>      <td>0.255047</td>      <td>1</td>    </tr>    <tr>      <th>2</th>      <td>0.201017</td>      <td>0</td>    </tr>    <tr>      <th>3</th>      <td>0.729307</td>      <td>1</td>    </tr>    <tr>      <th>4</th>      <td>0.148288</td>      <td>0</td>    </tr>  </tbody></table>

For some metrics, we will need to convert the probabilities into predicted class labels by setting a threshold. We will do it using the following function:

```python
# Run this cell, we will use the function in the exercises.

def threshold_probas(proba, threshold=.5): 
    if proba >= threshold:
        return 1
    else: 
        return 0 
```

## Exercise 1 - Understanding the problem with Accuracy 

In this exercise, you will implement a function to calculate the accuracy from true and predicted class labels. We will then compare the accuracy of predicted class labels for various predictions.

Implement below a function to compute the accuracy from a dataframe with true and predicted class labels:

- the `target` column contains the true labels
- the `prediction` column contains the model prediction for the class labels


<!-- grade_id: cell-72ba6d5cbf2be8e0, points: 0, grade: false, solution: true -->
```python
def accuracy(df):
    """
        Given a dataframe with `prediction` and `target` columns, 
        compute the accuracy metric
    
    Args:
        df (pd.DataFrame): the dataframe, with `prediction` and `target` columns

    Returns:
        accuracy (float): the accuracy metric    
    """    

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return accuracy

```

<!-- grade_id: cell-26325384c86158fa, points: 2, grade: true, solution: false -->
```python
df_1 = pd.DataFrame({"target": np.array([0, 0, 1, 1, 0, 1, 1, 1, 1, 0, 0])})

df_1_all_zeros = df_1.copy()
df_1_all_zeros["prediction"] = 0 
np.testing.assert_almost_equal(accuracy(df_1_all_zeros), 0.4545, 2)

df_1_all_ones = df_1.copy()
df_1_all_ones["prediction"] = 1
np.testing.assert_almost_equal(accuracy(df_1_all_ones), 0.54545, 2)

df_2 = pd.DataFrame({"target": np.array([0, 1, 1, 1, 0, 1, 1, 1, 1, 0, 1])})

df_2_all_zeros = df_2.copy()
df_2_all_zeros["prediction"] = 0 
np.testing.assert_almost_equal(accuracy(df_2_all_zeros), 0.2727, 2)

df_2_all_ones = df_2.copy()
df_2_all_ones["prediction"] = 1
np.testing.assert_almost_equal(accuracy(df_2_all_ones), 0.7272, 2)

```

Let's calculate the accuracy for predictions for a threshold = 0.5:

```python
# run this cell

df = load_data()
df["prediction"] = df['probas'].apply(threshold_probas)
print(f"Accuracy of predictions: {accuracy(df)}")
```

Now let's see what would happen if we just predicted 0 (majority class) or 1 for every data point:

```python
df["prediction"] = 0
print(f"Accuracy of predicting always zero: {accuracy(df)}")


df["prediction"] = 1
print(f"Accuracy of predicting always one: {accuracy(df)}")

```

As you see, the accuracy can be misleading. We got nearly the same accuracy for predicting all 0s and for our actual prediction with a threshold of 0.5 (0.84 vs. 0.85). If we optimized our model just for accuracy, we could end up with a model that just predicts all zero values. Let's move on to other metrics from the learning notebook.

We'll start with the confusion matrix and the insights it gives us.

## Exercise 2: Understanding metrics at a particular threshold 

We'll now look into confusion matrices. For this purpose we'll use a threshold of 0.3 and plot the confusion matrix. You will then use the numbers in the confusion matrix to calculate the true and false positives and negatives.

```python
# run this cell to calculate predictions for threshold = 0.3 and plot the corresponding confusion matrix

df = load_data()
df["prediction"] = df['probas'].apply(lambda prob: threshold_probas(prob, threshold=0.3))

plot_confusion_matrix(df["target"], df["prediction"])
```

### 2.1 FP, TP, FN, TN
Identify the following parameters for this confusion matrix and fill in the answers in the cell below:

- False Positives    (FP)
- True Positives     (TP)
- False Negatives    (FN)
- True Negatives     (TN)

<!-- grade_id: cell-bfc98777980f60ed, points: 0, grade: false, solution: true -->
```python
# true_positives
# TP = ...
# true_negatives 
# TN = ...
# false_positives
# FP = ...
# false_negatives
# FN = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f7dda939ca8ad6e4, points: 2, grade: true, solution: false -->
```python
assert hash_answer(TP) == 'a46e37632fa6ca51a13fe39a567b3c23b28c2f47d8af6be9bd63e030e214ba38'
assert hash_answer(TN) == '55f0124bb79f5c53d868ca45bbb0f4d04da15eea4fb29c6b95087fe8801bf0a3' 
assert hash_answer(FP) == '69f59c273b6e669ac32a6dd5e1b2cb63333d8b004f9696447aee2d422ce63763'
assert hash_answer(FN) == '98010bd9270f9b100b6214a21754fd33bdc8d41b2bc9f9dd16ff54d3c34ffd71'
```

Great job! We'll now use these parameters to compute more metrics. Implement below a function for each of these metrics: 

- False Positive Rate  (FPR)
- True Positive Rate (TPR)
- Precision
- Recall

Note that all functions take the same arguments although some of them might be superfluous in each case.

### 2.2 False Positive Rate

<!-- grade_id: cell-516897dd70339b47, points: 0, grade: false, solution: true -->
```python
def false_positive_rate(TP, TN, FP, FN):
    """
    Given the true positives, true negatives, false positives, and false negatives
    compute the false positive rate
    
    Args:
        TP (int): number of true positives     
        TN (int): number of true negatives     
        FP (int): number of false positives     
        FN (int): number of false positives     


    Returns:
        false_positive_rate (float): false positive rate     
    """    

    # false_positive_rate = ...  (a.k.a False Alarm Rate)
    # YOUR CODE HERE
    raise NotImplementedError()
    return false_positive_rate
```

<!-- grade_id: cell-d472dab689d7cb50, points: 1.5, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(false_positive_rate(10, 40, 13, 42), 0.2452, 2)
np.testing.assert_almost_equal(false_positive_rate(24, 42, 4, 64), 0.0869, 2)
np.testing.assert_almost_equal(false_positive_rate(64, 26, 64, 15), 0.7111, 2)
np.testing.assert_almost_equal(false_positive_rate(17, 15, 1, 3), 0.0625, 2)
```

### 2.3 True Positive Rate

<!-- grade_id: cell-a740e1c5ee69421a, points: 0, grade: false, solution: true -->
```python
def true_positive_rate(TP, TN, FP, FN):
    """
    Given the true positives, true negatives, false positives, and false negatives
    compute the true positive rate
    
    Args:
        TP (int): number of true positives     
        TN (int): number of true negatives     
        FP (int): number of false positives     
        FN (int): number of false positives     


    Returns:
        true_positive_rate (float): true positive rate     
    """    

    # true_positive_rate = ... 
    # YOUR CODE HERE
    raise NotImplementedError()
    return true_positive_rate
```

<!-- grade_id: cell-caa8049e86187fa1, points: 1.5, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(true_positive_rate(10, 40, 13, 42), 0.1923, 2)
np.testing.assert_almost_equal(true_positive_rate(24, 42, 4, 64), 0.2727, 2)
np.testing.assert_almost_equal(true_positive_rate(64, 26, 64, 15), 0.8101, 2)
np.testing.assert_almost_equal(true_positive_rate(17, 15, 1, 3), 0.8500, 2)    
```

### 2.4 Precision

<!-- grade_id: cell-c66a7234a25f7dc8, points: 0, grade: false, solution: true -->
```python
def precision_metric(TP, TN, FP, FN):
    """
    Given the true positives, true negatives, false positives, and falsenegatives
    compute the precision
    
    Args:
        TP (int): number of true positives     
        TN (int): number of true negatives     
        FP (int): number of false positives     
        FN (int): number of false positives     


    Returns:
        precision (float): precision score     
    """    
    # precision = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return precision
```

<!-- grade_id: cell-a10585e0d5167a4d, points: 1.5, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(precision_metric(10, 40, 13, 42), 0.4348, 2)
np.testing.assert_almost_equal(precision_metric(24, 42, 4, 64), 0.8571, 2)
np.testing.assert_almost_equal(precision_metric(64, 26, 64, 15), 0.5000, 2)
np.testing.assert_almost_equal(precision_metric(17, 15, 1, 3), 0.94444, 2)    
```

### 2.5 Recall

<!-- grade_id: cell-4cf0fcc674b2d58f, points: 0, grade: false, solution: true -->
```python
def recall_metric(TP, TN, FP, FN):
    """
    Given the true positives, true negatives, false positives, and false negatives
    compute the recall
    
    Args:
        TP (int): number of true positives     
        TN (int): number of true negatives     
        FP (int): number of false positives     
        FN (int): number of false positives     


    Returns:
        recall (float): recall score     
    """    
    # recall = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return recall
```

<!-- grade_id: cell-51b114e42c819d0f, points: 1.5, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(recall_metric(10, 40, 13, 42), 0.1923, 2)
np.testing.assert_almost_equal(recall_metric(24, 42, 4, 64), 0.2727, 2)
np.testing.assert_almost_equal(recall_metric(64, 26, 64, 15), 0.8101, 2)
np.testing.assert_almost_equal(recall_metric(17, 15, 1, 3), 0.8500, 2)    
```

Finally, let's apply our metrics to the dataframe we had. We'll compute these for the threshold of 0.3, which corresponds to the confusion matrix presented above.

```python
FPR = false_positive_rate(TP, TN, FP, FN)
TPR = true_positive_rate(TP, TN, FP, FN)
precision = precision_metric(TP, TN, FP, FN)
recall = recall_metric(TP, TN, FP, FN)

print('True Positives:       %0.0f' % TP)
print('True Negatives:       %0.0f' % TN)
print('False Positives:      %0.0f' % FP)
print('False Negatives:      %0.0f' % FN)
print('False Positive Rate:  %0.2f' % FPR)
print('True Positive Rate:   %0.2f' % TPR)
print('precision:            %0.2f' % precision)
print('recall:               %0.2f' % recall)

```

_Note: your `True Positive Rate` and `Recall` should be the same because... ehm... check their formula!_

## Exercise 3: model selection 

Consider the following two confusion matrixes: 

<img src="media/conf_mats.png" alt="drawing" width="600"/> 

Which model (A or B) is better if they represent the performance of a judge over their career? 
In this situation: 
- 1 is guilty
- 0 is innocent 
- predicting 1 sends people to jail
- predicting 0 sends people home free


<!-- grade_id: cell-b50fa7cfe6ed2f83, points: 0, grade: false, solution: true -->
```python
# best_option_for_a_judge = ... ("A" or "B")
# YOUR CODE HERE
raise NotImplementedError()
```

Which option (A or B) is best, if this is a model for cancer screening? 
In this situation: 
- 1 is cancer
- 0 is healthy
- predicting 1 sends people to a cancer screening
- predicting 0 sends people home 

<!-- grade_id: cell-67d2c8b099960bb4, points: 0, grade: false, solution: true -->
```python
# best_option_for_cancer_screening = ... ("A" or "B")
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f6e108311857c10, points: 1, grade: true, solution: false -->
```python
assert hash_answer(best_option_for_cancer_screening.lower()) == '3e23e8160039594a33894f6564e1b1348bbd7a0088d42c4acb73eeaed59c009d' 
assert hash_answer(best_option_for_a_judge.lower()) == 'ca978112ca1bbdcafac231b39a23dc4da786eff8147c4e72b9807785afee48bb'
```

## Exercise 4: Understanding the ROC Curve 

In this exercise, you will calculate the ROC curve step-by-step. You get the first two functions for free and you need to implement the rest.

```python
# functions for free

def threshold_probas(proba, threshold): 
    """
    Returns the prediction for a single observation based on a threshold
    
    Args:
        proba (float): a scores or probability, between 0.0 and 1.0
        threshold (float): a number between 0.0 and 1.0    
    
    Returns:
        the prediction: 1 or 0
    """ 
    if proba >= threshold:
        return 1
    else: 
        return 0 
    
def get_predictions(probas, threshold):
    """
    Returns the predictions for a series of observations
    
    Args:
        probas (pd.Series): a series of floats (scores or probabilities) between 0.0 and 1.0
        threshold (float): a number between 0.0 and 1.0    
    Returns:
        predictions (pd.Series): a series of 1s and 0s 
    """
    return probas.apply(lambda x: threshold_probas(x, threshold))
```

### 4.1. Confusion Matrix Elements 
A second take on the confusion matrix. Write a function that calculates all elements of the Confusion Matrix (TP, FP, TN, FN) from a dataset of true and predicted class labels.

<!-- grade_id: cell-0a058e0a8e2834d4, points: 0, grade: false, solution: true -->
```python
def compute_confusion_matrix(predictions, target):
    """
    Calculates all the elements of the confusion matrix
    
    Args:
        predictions (pd.Series): a series of ints, either 0s or 1s, with predictions
        target (pd.Series): a series of ints, either 0s or 1s, with observed outcomes 
    
    Returns:
        true_positives, true_negatives, false_positives, false_negatives (int):
                                                    elements of the confusion matrix
    """
    
    # it should take ~2 lines to calculate each measure
    # YOUR CODE HERE
    raise NotImplementedError()
    return true_positives, true_negatives, false_positives, false_negatives
```

<!-- grade_id: cell-2a5fde6327485279, points: 1, grade: true, solution: false -->
```python
df = load_data()
threshold_for_testing = 0.6
predictions = get_predictions(df['probas'], threshold_for_testing)
tp, tn, fp, fn = compute_confusion_matrix(predictions, df.target)

assert hash_answer(tp) == '4a44dc15364204a80fe80e9039455cc1608281820fe2b24f1e5233ade6af1dd5'
assert hash_answer(fp) == 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d'
assert hash_answer(tn) == 'd6723fa996ced47773f2dea29cce9b11f951e6dafe321a84ac7d32791c3b4660'
assert hash_answer(fn) == '2abaca4911e68fa9bfbf3482ee797fd5b9045b841fdff7253557c5fe15de6477'
```

### 4.2. True/False Positives Rates
Write a function that calculates the True Positives Rate and the False Positives Rate from the true and predicted class labels. You can use the functions you have implemented above.

<!-- grade_id: cell-876416c5941a720c, points: 0, grade: false, solution: true -->
```python
def calculate_true_and_false_positives_rate(predictions, target):
    """
    Calculates the true and false positives rate
    
    Args:
        predictions (pd.Series): a series of ints, either 0s or 1s, with predictions
        target (pd.Series): a series of ints, either 0s or 1s, with observed outcomes 
    
    Returns:
        true_positives_rate, false_positives_rate (float): true and false positives rate
    """   
    # YOUR CODE HERE
    raise NotImplementedError()
    return true_positives_rate, false_positives_rate
```

<!-- grade_id: cell-c4517396cfce00b0, points: 1, grade: true, solution: false -->
```python
df = load_data()
threshold_for_testing = 0.6
predictions = get_predictions(df['probas'], threshold_for_testing)
tpr, fpr = calculate_true_and_false_positives_rate(predictions, df['target'])

assert hash_answer(round(tpr, 4)) == '918ae0d325e428351045f50747561dd25451e3866462a38c0e62ec2913fdd1a5'
assert hash_answer(round(fpr, 4)) == '4c158422ed9316829516c182728418409140930df77cf52b47ea03cb59bfd2a3'
```

### 4.3. ROC Curve
Now you have all the functions that help you to calculate the ROC curve. Implement a function that calculates the true and false positives rates for a list of thresholds. The function should return two dictionaries, for the true/false positives rate, where the keys are the thresholds and the values are the calculated true/false positives rates.

<!-- grade_id: cell-845d59b49ea2fe9f, points: 0, grade: false, solution: true -->
```python
def hand_made_roc_curve(probas, target, list_of_thresholds):
    """
    Calculates the true and false positives rates for the given thresholds.
    
    Args:
        probas (float): a list of probabilities for some observations, between 0.0 and 1.0
        target (int): a list of true class labels for the same observations, 0s and 1s
        list_of_thresholds (float): a list of thresholds for which to calculate the true
                                    and false positive rate
    Returns:
        TPRs, FPRSs: dictionaries with true and false positives rates for each threshold,
                     keys - thresholds, values - corresponding tpr/fpr
        
    """
    
    TPRs = {}
    FPRs = {}
    
    # for each threshold, get the predictions and calculate the TPR and FPR
    # you already defined everything above, so each line should be just a function call   
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return TPRs, FPRs
```

<!-- grade_id: cell-5d25008fa21ed0ab, points: 1, grade: true, solution: false -->
```python
df = load_data()
list_of_thresholds = np.linspace(0, 10) ** 2 / 100

hand_made_tpr, hand_made_fpr = hand_made_roc_curve(df['probas'], df['target'], list_of_thresholds)

assert isinstance(hand_made_tpr, dict), "TPRs is not a dictionary"
assert isinstance(hand_made_fpr, dict), "FPRs is not a dictionary"

assert sum(list(hand_made_tpr.keys())== list_of_thresholds) == len(list_of_thresholds), 'The keys of the TPR dictionary are not correct'
assert sum(list(hand_made_fpr.keys())== list_of_thresholds) == len(list_of_thresholds), 'The keys of the FPR dictionary are not correct'

tpr_array = np.array(list(hand_made_tpr.values()))
fpr_array = np.array(list(hand_made_fpr.values()))

np.testing.assert_almost_equal(tpr_array.mean(), 0.56, 2, err_msg='The values of the TPR dictionary are not correct.')
np.testing.assert_almost_equal(tpr_array.var(), 0.14, 2, err_msg='The values of the TPR dictionary are not correct.')
np.testing.assert_almost_equal(tpr_array.min(), 0., 2, err_msg='The values of the TPR dictionary are not correct.')
np.testing.assert_almost_equal(tpr_array.max(), 1., 2, err_msg='The values of the TPR dictionary are not correct.')

np.testing.assert_almost_equal(fpr_array.mean(), 0.29, 2, err_msg='The values of the FPR dictionary are not correct.')
np.testing.assert_almost_equal(fpr_array.var(), 0.10, 2, err_msg='The values of the FPR dictionary are not correct.')
np.testing.assert_almost_equal(fpr_array.min(), 0.0, 1, err_msg='The values of the FPR dictionary are not correct.')
np.testing.assert_almost_equal(fpr_array.max(), 1.0, 1, err_msg='The values of the FPR dictionary are not correct.')
```

Finally, the fruits of your labor. We will plot an ROC curve calculated with your function:

```python
df = load_data()

# generating a space of thresholds (we'll be more "dense" near zero)
list_of_thresholds = np.linspace(0, 10) ** 2 / 100

# using the function you have just created 
hand_made_tpr, hand_made_fpr = hand_made_roc_curve(df['probas'], df['target'], list_of_thresholds)

# just making a dataframe to pretty things up
hand_made_roc_df = pd.DataFrame({'False Positive Rate': hand_made_fpr, 
                        'True Positive Rate': hand_made_tpr})

# naming the index
hand_made_roc_df.index.name = 'Thresholds'

# plotting the ROC curve 
hand_made_roc_df.set_index('False Positive Rate').plot();
plt.show()

# displaying the ROC curve data
display(hand_made_roc_df.iloc[10:20])
```

## Exercise 5: The joy of scikit 
Ok, you earned it - you get to use scikit now ! :)

Run the cell below to load the data and create predicted class labels.

```python
# Run this cell

df = load_data()

df['prediction'] = get_predictions(df['probas'], threshold=.3)
df.head()
```

### 5.1. Scikit Accuracy/Precision/Recall
Implement a function to calculate the accuracy, precision, and recall with sklearn. Don't forget to import the necessary functions! And don't confuse the inputs. :) 

<!-- grade_id: cell-67c2e8ef543f95ff, points: 0, grade: false, solution: true -->
```python
def accuracy_precision_recall_sklearn(predictions, target):
    """
    Given the predicted and true labels, return the accuracy, precision,
    and recall using sklearn
    
    Args:
        predictions (pd.Series or np.array): classifier class predictions
        target (pd.Series or np.array): true labels     

    Returns:
        accuracy, precision, recall (float): metrics calculated with sklearn
    """    
    # NOTE: Even though this is not a good programming practice, just 
    #       for the sake of the exercise, ensure you import the right 
    #       functions within this code:
    # from ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return accuracy, precision, recall
```

<!-- grade_id: cell-78d7b23cb6126162, points: 2, grade: true, solution: false -->
```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)

df_1 = df.iloc[100:200].copy()
accuracy, precision, recall = accuracy_precision_recall_sklearn(df_1['prediction'], df_1['target'])
np.testing.assert_almost_equal(accuracy, 0.83, 2, err_msg="The accuracy is not correct.")
np.testing.assert_almost_equal(precision, 0.5714, 2, err_msg="The precision is not correct.")
np.testing.assert_almost_equal(recall, 0.7619, 2, err_msg="The recall is not correct.")

df_2 = df.iloc[0:50].copy()
accuracy, precision, recall = accuracy_precision_recall_sklearn(df_2['prediction'], df_2['target'])
np.testing.assert_almost_equal(accuracy, 0.86, 2, err_msg="The accuracy is not correct.")
np.testing.assert_almost_equal(precision, 0.5, 2, err_msg="The precision is not correct.")
np.testing.assert_almost_equal(recall, 0.5714, 2, err_msg="The recall is not correct.")

```

Let's check the metrics for the full dataset:

```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)
accuracy, precision, recall = accuracy_precision_recall_sklearn(df['prediction'], df['target'])
print(f"Accuracy: {accuracy}")
print(f"Precision: {precision}")
print(f"Recall: {recall}")
```

Notice the great accuracy (which, as we know, is just trying to delude us), but not so great precision and OKish recall.

Next level - confusion matrix with sklearn.

### 5.2. Scikit Confusion Matrix
Implement a function that return the confusion matrix for the given true and predicted class labels. Use sklearn and don't forget to import. Don't confuse the function inputs!

<!-- grade_id: cell-86d92e523ba4389c, points: 0, grade: false, solution: true -->
```python
def confusion_matrix_sklearn(predictions, target):
    """
    Given the predicted and true labels, return the confusion
    matrix using scikit-learn
    
    Args:
        predictions (pd.Series or np.array): classifier predictions     
        target (pd.Series or np.array): true labels     

    Returns:
        conf_mat: confusion_matrix
    """    
    # from ? import ...
    # conf_mat = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return conf_mat
```

<!-- grade_id: cell-4e5021ddb84188d6, points: 2, grade: true, solution: false -->
```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)

df_1 = df.iloc[100:200].copy()

np.testing.assert_array_almost_equal(
    confusion_matrix_sklearn(df_1['prediction'], df_1['target']), 
    [[67, 12], [5,  16]], 
    2
)

df_2 = df.iloc[0:50].copy()

np.testing.assert_array_almost_equal(
    confusion_matrix_sklearn(df_2['prediction'], df_2['target']), 
    [[39, 4], [3, 4]], 
    2
)
```

And we can check the matrix for the full dataset:

```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)

print(f"Confusion matrix:\n {confusion_matrix_sklearn(df['prediction'], df['target'])}")
print("Meaning:\n [[TP   FN]\n [ FP  TN]]")
```

We're pretty good in identifying positives, but quite bad in detecting negatives.

### 5.3. Scikit ROC Curve
Write a function which calculates the ROC curve and AUROC score using sklearn. The function should return the list of false positives rates, true positives rates, and the thresholds, and the AUROC score. As usual, don't forget the imports and don't confuse the inputs.

<!-- grade_id: cell-8a1d01de0d2450f8, points: 0, grade: false, solution: true -->
```python
def get_roc_sklearn(probas, target):
    """
    Given the predicted probabilities (probas) and true labels, 
    return the information regarding the roc curve - false positives rate,
    true positives rate and thresholds - and the area under the curve
    
    Args:
        probas (pd.Series or np.array): classifier probabilities     
        target (pd.Series or np.array): true labels     

    Returns:
        fpr: false positives rates array for roc curve
        tpr: true positives rates array for roc curve
        thresholds: thresholds array for roc curve
        roc_auc: roc area under the curve

    """    

    # from ? import ... 
    # YOUR CODE HERE
    raise NotImplementedError()
    return fpr, tpr, thresholds, roc_auc

```

<!-- grade_id: cell-74e565ebc3dd0dc9, points: 2, grade: true, solution: false -->
```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)

df_1 = df.iloc[100:200].copy()
fpr_1, tpr_1, thresholds_1, roc_auc_1 = get_roc_sklearn(df_1['probas'], df_1['target'])

np.testing.assert_array_almost_equal(fpr_1[3:6], [0.013, 0.013, 0.025], 3)
np.testing.assert_array_almost_equal(tpr_1[6:8], [0.333, 0.333], 3)
np.testing.assert_array_almost_equal(thresholds_1[7:10], [0.502, 0.498, 0.477], 3)
np.testing.assert_almost_equal(roc_auc_1, 0.86197, 2)

df_2 = df.iloc[0:50].copy()
fpr_2, tpr_2, thresholds_2, roc_auc_2 = get_roc_sklearn(df_2['probas'], df_2['target'])

np.testing.assert_array_almost_equal(fpr_2[3:6], [0.047, 0.07 , 0.07], 3)
np.testing.assert_array_almost_equal(tpr_2[1:3], [0.143, 0.143], 3)
np.testing.assert_array_almost_equal(thresholds_2[7:10], [0.052, 0.014, 0.011], 3)
np.testing.assert_almost_equal(roc_auc_2, 0.7575, 2)
```

Finally, let's run it for the whole dataset and plot the curve:

```python
df = load_data()
df['prediction'] = get_predictions(df['probas'], threshold=.3)

fpr_test, tpr_test, thresholds_test, roc_auc_test = get_roc_sklearn(df['probas'], df['target'])
plt.figure(figsize=(4,4))
lw = 2
plt.plot(fpr_test, tpr_test, color='orange', lw=lw, label='ROC curve (AUROC = %0.2f)' % roc_auc_test)
plt.plot([0, 1], [0, 1], color='navy', lw=lw, linestyle='--', label='random')
plt.xlim([-0.05, 1.0])
plt.ylim([0.0, 1.05])
plt.grid()
plt.xlabel('False Positive Rate')
plt.ylabel('True Positive Rate')
plt.title('Receiver operating characteristic example')
plt.legend(loc="lower right")
plt.show()
```

Congratulations! You made it to the end of another unit! You are now ready to explore more about classification and regression problems and other particularities of model training and selection.

See you in the next unit!