# BLU02 - Exercises Notebook

```python
import hashlib # for grading

import os
import pandas as pd
import numpy as np
import datetime
from sklearn.model_selection import train_test_split
from sklearn.base import BaseEstimator, TransformerMixin
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler, RobustScaler
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_absolute_error
```

## 1 Read the Oscars data (graded)

In this first exercise, we aim to create a single dataframe, combining all Oscar nominees from all ceremonies.

With a caveat though: **we want to include seasons from the year 1960 onwards**.

<!-- grade_id: cell-e16abeb47fc4ea37, points: 0, grade: false, solution: true -->
```python
def read_year(folder_path, file_name):
    path = os.path.join(folder_path, file_name)
    return pd.read_csv(path, index_col = 0)

def read_nominees(folder_path):
    files = os.listdir(folder_path)
    # Create a list with the name of all files containing annual nominees from
    # 1960 inclusive and onwards (just the filename, no complete path.)
    # files_from_1960: List[str] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Create a list with the dataframes
    # nominees_year: List[pd.DataFrame] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Use pd.concat to create a single dataframe.
    # nominees: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Drop the column 'ceremony'.
    # nominees = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove missing values.
    # nominees = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return nominees

nominees = read_nominees(os.path.join('data', 'oscars'))
```

<!-- grade_id: cell-2656708135a5a5e4, points: 2, grade: true, solution: false -->
```python
assert nominees['year_ceremony'].min() == 1960
assert nominees['year_ceremony'].max() == 2023
assert nominees.isna().sum().sum() == 0
assert nominees.shape == (7116, 6)
```

## 2 Read the IMDB Ratings data (graded)

<!-- grade_id: cell-e228c2d82463c6b4, points: 0, grade: false, solution: true -->
```python
def read_ratings(file_path): 
    # Read the ratings data and drop the 'director', 'star1', 'star2', 'star3', and 'star4' columns.
    # top_rated: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Please make the necessary changes and convert the 'runtime' column to int
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove the lines with no metascore info.
    # YOUR CODE HERE
    raise NotImplementedError()
    return top_rated

top_rated = read_ratings(os.path.join('data','imdb_top_1000.csv'))
```

<!-- grade_id: cell-3d21007e725ab889, points: 2, grade: true, solution: false -->
```python
assert top_rated.shape == (843, 5)
assert top_rated.runtime.min() == 64
assert top_rated.runtime.max() == 321
assert top_rated.metascore.isna().sum() == 0
assert set(top_rated.columns) == set([
    'film', 'metascore', 'no_votes', 'rating', 'runtime'
])
```

## 3 Combine Oscars and Ratings data (graded)

Let's combine both dataframes into a single dataset, using an inner join.

<!-- grade_id: cell-a65f1464b4525a8b, points: 0, grade: false, solution: true -->
```python
# Remember that you want use a column of both dataframes to combine them.
# Join only the nominees of films present on the ratings list
# best_rated_nominees = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ac9aef3d5251e36c, points: 2, grade: true, solution: false -->
```python
assert best_rated_nominees.shape == (1919, 10)
assert set(best_rated_nominees.columns) == set(['year_film', 'year_ceremony', 'category', 'name', 'film', 'winner',
       'runtime', 'rating', 'metascore', 'no_votes'])
```

## 4 Read top grossing and budget films data (graded)

We will read the two remaining pieces of data. 

Again, albeit the step-by-step description, we encourage you to use method chaining.

<!-- grade_id: cell-503e208490ff38e0, points: 0, grade: false, solution: true -->
```python
def read_gross(file_path):
    # Read the works data.
    # top_grossing: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Remove the year column.
    # gross: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Ensure that the gross data is read as int.
     # top_grossing: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return top_grossing


def read_budget(file_path):
    # Read the top budget data and drop the 'runtime', 'theaters', and 'year' Columns
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Please make the necessary changes and convert the 'budget_rank' column to int
    # YOUR CODE HERE
    raise NotImplementedError()
    return top_budget


top_grossing = read_gross('data/gross_top_200.csv')
top_budget = read_budget('data/budget_top_500.csv')
```

<!-- grade_id: cell-b8389314995f18ea, points: 2, grade: true, solution: false -->
```python
assert top_grossing.shape == (200, 3)
assert set(top_grossing.columns) == set([
    'gross_rank', 'film', 'gross'
])

assert top_budget.shape == (500, 3)
assert set(top_budget.columns) == set([
   'budget_rank', 'film', 'production_cost'
])
```

## 5 Combine the top grossing and budget films

Like we did for Oscar nominees and top rated films, now we combine the top budget and grossing films.

<!-- grade_id: cell-08d9a086cc5646cf, points: 0, grade: false, solution: true -->
```python
# Combine both dataframes, again using an inner type of join
# top_grossing_budget : pd.DataFrame = ....
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d9f103dfffd311b, points: 2, grade: true, solution: false -->
```python
assert top_grossing_budget.shape == (129, 5)
assert set(top_grossing_budget.columns) == set(
    [
        'budget_rank', 'film', 'production_cost', 'gross_rank', 'gross'
    ]
)
```

## 6 Combine everything (graded)

The final goal here is to create a single dataframe.

<!-- grade_id: cell-ce1f05022e8cd63a, points: 0, grade: false, solution: true -->
```python
# Combine best_rated_nominees and nyp into a single dataframe.
# You need to figure out the common column shared between the two dataframes
# top_films = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce29fa3aec1c244e, points: 2, grade: true, solution: false -->
```python
assert top_films.shape == (184, 14)
assert set(top_films.columns) == set(
    [
       'year_film', 'year_ceremony', 'category', 'name', 'film', 'winner',
       'runtime', 'rating', 'metascore', 'no_votes', 'budget_rank',
       'production_cost', 'gross_rank', 'gross'
    ]
)
```

## 7 Final transformations (graded)

Now, we perform the train-test split.

We also perform some final transformations on both datasets:

* Tranform "winner" into a binary feature
* Create a new feature, rating_rank, from the rating column. 
* Filter out the movies that appear less than 10 times in the DataFrame.
* Keep only 'film', 'winner', 'category', 'runtime', 'rating', 'metascore', 'no_votes', 'budget_rank', 'production_cost', 'gross_rank', 'rating_rank' and 'gross' columns.

<!-- grade_id: cell-1c1ab0d912e615ca, points: 0, grade: false, solution: true -->
```python
def transform_winner(df):
    df = df.copy()
    # df = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    return df

def create_rating_rank(df):
    df = df.copy()
    # df['rating_rank'] = 
    # YOUR CODE HERE
    raise NotImplementedError()
    return df

def preprocess_data(df):
    # You should follow these exact steps:
    #   1 - Binarize 'winner'
    #   2 - Filter out rows that have a film that appear is less than 10 times in the DataFrame.
    #   3 - Keep only the 'film', 'winner', category', 'runtime', 'rating', 'metascore', 'no_votes', 
    #      'budget_rank', 'production_cost', 'gross_rank', 'rating_rank' and 'gross' columns
    #   5 - Create a new feature, rating_rank, from the ranking column. 
    #   6 - Sort the DataFrame by 'rating' in ascending order
    # YOUR CODE HERE
    raise NotImplementedError()
    return df


top_films_preprocessed = preprocess_data(top_films)
```

<!-- grade_id: cell-9d9c75c48e4eaf63, points: 2, grade: true, solution: false -->
```python
assert top_films_preprocessed.shape == (59, 12)
assert set(top_films_preprocessed.columns) == {
       'category', 'film', 'runtime', 'rating', 'metascore', 'no_votes',
       'budget_rank', 'production_cost', 'gross_rank', 'gross', 'rating_rank','winner'
}
assert top_films_preprocessed.budget_rank.min() == 49
assert top_films_preprocessed.gross_rank.max() == 180
assert top_films_preprocessed.iloc[0].no_votes == 769145
```

# The house prices dataset

A dataset containing several characteristics of several houses and their selling price 

* LotFrontage: Linear feet of street connected to property
* LotArea: Lot size in square feet
* OverallQual: Rates the overall material and finish of the house
       10  Very Excellent
       9	Excellent
       8	Very Good
       7	Good
       6	Above Average
       5	Average
       4	Below Average
       3	Fair
       2	Poor
       1	Very Poor
* OverallCond: Rates the overall condition of the house

       10	Very Excellent
       9	Excellent
       8	Very Good
       7	Good
       6	Above Average	
       5	Average
       4	Below Average	
       3	Fair
       2	Poor
       1	Very Poor
* MasVnrArea: Masonry veneer area in square feet
* BsmtFinSF1: Type 1 finished square feet
* BsmtUnfSF: Unfinished square feet of basement area
* TotalBsmtSF: Total square feet of basement area
* 1stFlrSF: First Floor square feet
* 2ndFlrSF: Second floor square feet
* LowQualFinSF: Low quality finished square feet (all floors)
* GrLivArea: Above grade (ground) living area square feet
* BsmtFullBath: Basement full bathrooms
* BsmtHalfBath: Basement half bathrooms
* FullBath: Full bathrooms above grade
* HalfBath: Half baths above grade
* BedroomAbvGr: Bedrooms above grade (does NOT include basement bedrooms)
* KitchenAbvGr: Kitchens above grade
* TotRmsAbvGrd: Total rooms above grade (does not include bathrooms)
* Fireplaces: Number of fireplaces
* GarageCars: Size of garage in car capacity
* GarageArea: Size of garage in square feet
* WoodDeckSF: Wood deck area in square feet
* OpenPorchSF: Open porch area in square feet
* EnclosedPorch: Enclosed porch area in square feet
* 3SsnPorch: Three season porch area in square feet
* ScreenPorch: Screen porch area in square feet
* PoolArea: Pool area in square feet
* MiscVal: $Value of miscellaneous feature 
* SellingDate: Date when the house was sold
* BuildingDate: Date when the house was built
* RemodAddDate: Remodel date (same as construction date if no remodeling or additions)
* SalePrice: The house price at the selling date (our target variable)

Let's read the csv and create our train-test-split

```python
def house_price_dataset():
    return pd.read_csv(
    'data/housePrices.csv', 
        parse_dates=[
            'SellingDate',
            'BuildingDate',
            'RemodAddDate'
        ]
    )

dataset = house_price_dataset()
dataset_train, dataset_test = train_test_split(dataset, random_state=0)
X_train = dataset_train.drop(columns='SalePrice')
y_train = dataset_train.SalePrice
X_test = dataset_test.drop(columns='SalePrice')
y_test = dataset_test.SalePrice
```

## 8 Build a DateTransformer transformer (graded)

A simple transformer that transforms dates into timedeltas can be useful, from times to times, when modeling. Usually when you have features that are Dates you compute a time delta between the feature and a given refence date.

e.g Imagine that your clients have a loyalty period that ends at a given date. When your model is doing some predictions, one of the features that you can use is the number of days until the end of the loyalty period. i.e the date when the loyalty ends minus the date when your model is running. 

In the house prices dataset, the selling date will be the reference data, since we want to predict the house price at the selling date. For instance, two houses with the exact same features can vary in prices if the construction year is different. So we should input this information and feed into the model. Then we need to convert the other dates using our transformer

Hint: Result should be integers

<!-- grade_id: cell-d9538038181ba578, points: 0, grade: false, solution: true -->
```python
class DateTransformer(BaseEstimator, TransformerMixin):
    # Implement the __init__ method.
    # Our DateTransformer must be able to receive two parameters: 
    # datetime_cols: a list, that contains the datetime cols that should be converted
    # ref_date_col: indicates the name of the column that should be used as reference date,
    # YOUR CODE HERE
    raise NotImplementedError()
        
    # There's no need for a fit method in this case, it does nothing.
    # We should be able to call fit without any explicit parameters.
    # Meaning: we should be able to call fit() on the transformer.
    # YOUR CODE HERE
    raise NotImplementedError()

    # Transform should transform all datetime columns into the difference in days to the reference date.
    # The reference date column should be dropped. 
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b64b26753ecd8561, points: 2, grade: true, solution: false -->
```python
X_train_transformed = DateTransformer(
    datetime_cols=['BuildingDate', 'RemodAddDate'], 
    ref_date_col='SellingDate'
).fit_transform(X_train)
assert X_train_transformed.BuildingDate.min() == -49008
assert X_train_transformed.BuildingDate.max() == -1
assert 'SellingDate' not in X_train_transformed.columns
assert X_train_transformed.dtypes.BuildingDate == np.dtype('int64')
assert X_train_transformed.dtypes.RemodAddDate == np.dtype('int64')
```

You might be wondering why we have to implement it as a Transformer Class, and not using functions.
You'll understand the reason in the next section, we're we can tie them all together in a `Pipeline`.

## 9 Building the pipeline (graded)

Finally, we want to use the two transformers together and run a linear regression on top. We want to:

* Convert the dates to time deltas relative to the Selling Date.

* Scale all features to the same range, using `sklearn.preprocessing.RobustScaler()`.

* Estimate the SellingPrice using a Linear Regression.

Standardization of datasets is a common requirement for many machine learning estimators implemented in scikit-learn; they might behave badly if the individual features do not more or less look like standard normally distributed data (i.e., Gaussian with zero mean and unit variance).

In practice we often ignore the shape of the distribution and just transform the data to center it by removing the mean value of each feature, then scale it by dividing non-constant features by their standard deviation.

For instance, many elements used in the objective function of a learning algorithm (such as the RBF kernel of Support Vector Machines or the l1 and l2 regularizers of linear models) assume that all features are centered around zero and have variance in the same order. If a feature has a variance that is orders of magnitude larger than others, it might dominate the objective function and make the estimator unable to learn from other features correctly as expected.



```python
X_train.describe()
```

<!-- grade_id: cell-6bc09de5e71383d3, points: 0, grade: false, solution: true -->
```python
# Create a pipeline including:
#   1 - 'date_converter', DateTransformer(['BuildingDate', 'RemodAddDate'], ref_date_col='SellingDate')
#   2 - 'robust_scaler', RobustScaler() with the default parameters
#   3 - 'model', LinearRegression
# YOUR CODE HERE
raise NotImplementedError()


pipeline.fit(X_train, y_train)
y_pred = pipeline.predict(X_test)

mae = mean_absolute_error(y_test, y_pred)
print('MAE: {}'.format(mae))
```

<!-- grade_id: cell-5d21dcdd34a13d24, points: 4, grade: true, solution: false -->
```python
assert type(pipeline) == Pipeline
assert type(pipeline.named_steps['date_converter']) == DateTransformer
assert type(pipeline.named_steps['robust_scaler']) == RobustScaler
assert pipeline.named_steps['date_converter'].get_params()['ref_date_col'] == 'SellingDate'
assert set(
    pipeline.named_steps['date_converter'].get_params()['datetime_cols']
) == {'BuildingDate', 'RemodAddDate'}
assert type(pipeline.named_steps['model']) == LinearRegression 
```

## 10. Access the cofficients from the pipeline (ungraded)

Now we would want to obtain the coefficients from the model to understand features with the most predictive power.

<!-- grade_id: cell-58083ef8ab8d97c7, points: 0, grade: false, solution: true -->
```python
#coefs = ....
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f42cbedb988b6aa8, points: 0, grade: true, solution: false -->
```python
assert coefs.shape == (30,), 'Wrong number of coefficients. Did you select the features correctly?'
```

Exercises complete, congratulations! You are about to become a certified data wrangler.